```python
# @hidden_cell
""" Bart Gerritsen, Oct 2018:

Note: for safety and robustness, styles and script are contained inside
      the Notebook rather than in external *.css and *.js files
"""      
from IPython.core.display import HTML
from IPython.display import display
tag = HTML('''
<style>
    /*TU color table */
    :root {
      --tu-black:        rgb(0,0,0);
      --tu-white:        rgb(255,255,255);
      --tu-cyan:         rgb(0,166,214);
      --tu-green:        rgb(165,202,26);
      --tu-yellow:       rgb(225,196,0);
      --tu-orange:       rgb(230,70,22);
      --tu-red:          rgb(225,26,26);
      --tu-purple:       rgb(109,23,127);
      --tu-slategreen:   rgb(107,134,137);
      --tu-turqoise:     rgb(0,136,145);
      --tu-darkblue:     rgb(29,28,115);
      --tu-skyblue:      rgb(110,187,213);
    }
    h2, h3, h4 {
        background-color: var(--tu-white);
        color: var(--tu-cyan);
    }
    h1 {
        background-color: var(--tu-cyan);
        color: var(--tu-white);
    }
    em {
        color: var(--tu-cyan);
    }
     
    div.output_stdout {
        background-color: var(--tu-green);
        color: var(--tu-black);
    }
    div.output_stdout:before {
        content: "stdout output;";
    }
    div.output_stderr {
        background-color: var(--tu-yellow);
        color: var(--tu-black);
    }
    div.output_stderr:before {
        content: "stderr output;";
    }
</style>
<script>
    code_show=true; 
    IPython.OutputArea.prototype._should_scroll = function(lines) {
        return false;
    }
    function code_toggle() {
        if (code_show){
            $('div.cell.code_cell.rendered.selected div.input').hide();
        } else {
            $('div.cell.code_cell.rendered.selected div.input').show();
        }
        code_show = !code_show
    }     
    $( document ).ready(code_toggle);
</script>
<a href="javascript:code_toggle()"><h4>Notebook settings</h4></a>
''')
display(tag)
```

Output:
```
<IPython.core.display.HTML object>
```

<header>
    <div style="overflow: auto;">
        <img src="./figures/TUDelft.jpg" style="float: left;" />
        <img src="./figures/DUT_Flame.png" style="float: right; width: 100px;" />
    </div>
    <div style="text-align: center;">
        <h2>Assignment A2: Binary Search Tree, Heap and Priority Queue (Python3)</h2>
        <h6>&copy; 2019, Bart Gerritsen, TU Delft.</h6>
    </div>
    <br>
    <br>
</header>

# Introduction <a class="anchor" id="introduction"/>
Below, you find **Assignment A2**, a 2-student team assignment, for which the team has exactly 2 weeks. For this assignment, 80 points can be scored altogether, assigning you a grade 8. During the assessment of you work, all team members should be able to explain, demo the work handed in, and answer background questions. During the assessment, you can raise your grade to a 10. On the other hand, assessors may lower the grade or even disqualify your work, in case your understanding of your own work appears poor or insufficient. In case of fraud, assessors will report this to the course responsible, without assinging a grade at all.    

Do **not** Run-all ```|>|>``` , as there are some long tasks in this Notebook. Start at the top and run cell-by-cell, advancing downwards.

## Overview of this assignment

The purpose of this assignment is to make you familiar with implementing a data structure in Python in an object oriented way. During the lectures, you were presented pseudo code of different basic data structures. Now we expect you to implement one of these structures yourself.

To make it clear what is needed, we will provide you with the following (skeletons of) classes: 

1. `class BSTNode` 
2. `class Heap`
3. `class PriorityQueue`
4. `class ServiceTicket`
5. `class PriorityQueueWithHeap`

These classes are just templates (skeletons), with incompletely implemented code, empty methods, etc. and your assignment is to come up with an implementation of these parts, so that the questions asked and the tasks assigned can be answered and demonstrated with running code.

# Exercise 1: implement a Binary Search Tree
An $N$-node binary search tree is a binary tree in which nodes are ordered, thus providing $\mathcal{O}(log~N)$ time complexity searching on average. Below is a template to implement a binary search tree using `class BSTNode`. In `class BSTNode`, every instance has `left`, a `right`, and a `data` field. 

To make you assignment easier, we provide you with an support `class BSTPrinter`, that can print the structure or the nodes of a BST in text, like demonstrated below. What the printer will print is the left diagram; you read this diagram as if it were the right disagram, with lines and boxes. You can call the *BSTPrinter* whenever you want, providing it just with a BST and an optional header you want to see printed.

On the other hand, you are expected to understand the concept of a static method (see [here](https://www.journaldev.com/18722/python-static-method) or [here](https://pythonbasics.org/static-method/)) before you start with this assignment.

<img width="760" src="./figures/bst_tree_print-int2.png"/>

<!-- grade_id: q_1_0_1, points: 0, grade: false, solution: true -->
```python
class BSTNode:
    
    @staticmethod
    def find(root, parent, key):
        """return node with key given and its parent, or None if not found"""
        ### BEGIN SOLUTION
        # key in the root?
        if root is None or root.data == key:
            return root, parent
        # continue searching down the tree ...
        # key greater than root's key?
        if key > root.data:
            return BSTNode.find(root.right, root, key)
        # Key is smaller than root's key
        return BSTNode.find(root.left, root, key)
        ### END SOLUTION
    
    @staticmethod
    def height(tree):
        
        ### BEGIN SOLUTION
        
        tree_height = 0
        if tree is not None:
            height_left  = BSTNode.height(tree.left)
            height_right = BSTNode.height(tree.right)
            tree_height  =max(height_left, height_right) + 1
        
        ### END SOLUTION
        
        return tree_height
    
    @staticmethod    
    def print_BSTNode(BSTNode):
        """visitor function printing nodal data"""
        print(f'{str(BSTNode):s}', end= ' ')

    # --------------------------
      
    def __init__(self, data):
        self.data = data
        self.left = None
        self.right = None

    def __str__(self):
        """return nodal data"""
        return self.data
    
    
    def child_count(self):
        """return nr of chldren"""
        count = 0
        if self.left is not None:
            count += 1
        if self.right is not None:
            count += 1
        return count
 

    def insert(self, data):
        """insert data in binary tree"""
        
        ### BEGIN SOLUTION
        
        if data > self.data:
            if self.right is None:
                self.right = BSTNode(data)
            else:
                # insert it in the left subtree ...
                self.right.insert(data)
        else:
            if self.left is None:
                self.left = BSTNode(data)
            else:
                # insert it in the right subtree ...
                self.left.insert(data)
        
        ### END SOLUTION

        
    def delete(self, key):
        """delete BSTNode with key given. Return T|F succesfully deleted"""
        
        # aux functions for delete ...
        def _delete_leaf(node, parent):
            """delete a leaf node with 0 children"""
            result = False
            
            ### BEGIN SOLUTION
            
            if parent is None:
                # it is the root node... 
                # ... no parent to update ...
                node = None
            else:
                # it is a true leaf node ...
                # ... update the correct link info ...
                if parent.left is node:
                    parent.left = None
                else:
                    parent.right = None
            del node
            result = True
            
            ### END SOLUTION
            
            return result
        
        def _delete_one(node, parent):
            """delete a node with 1 child. Return T|F succesfully deleted"""
            result = False
            
            ### BEGIN SOLUTION
            
            # ... shift up child to position of the node ...
            # ... connect node->child to parent ...
            if node.left is not None:
                toShiftUp = node.left
            else:
                toShiftUp = node.right
            if parent is None:
                # it is the root node ...
                # ... just copy ...
                node.left = toShiftUp.left
                node.right = toShiftUp.right
                node.set_data = toShiftUp.data
            else:
                # no, it's not the root ...
                if parent.left is node:
                    parent.left = toShiftUp
                else:
                    parent.right = toShiftUp
                del node
            result = True
            
            ### END SOLUTION
            
            return result       
        
        def _delete_two(node, parent):
            """delete a node with two subtrees. Return T|F succesfully deleted"""
            
            def _find_successor(node, parent):
                """return successor and parent of successor"""
                
                ### BEGIN SOLUTION
                
                parent_succ = node
                succ = node.right                
                while succ.left is not None:
                    parent_succ = succ
                    succ = succ.left
                # we reached the leaf ...
                
                ### END SOLUTION
                
                return succ, parent_succ
            
            result = False
            
            # we are always finding a successor ...
            succ, parent_succ = _find_successor(node, parent)
            # now copy the data ...
            
            ### BEGIN SOLUTION
            
            node.data = succ.data
            # ... and recursively remove my duplicate (succ)...
            # ... which can atmost have 1 child (to the right)
            if succ.right is not None:
                result = _delete_one(succ, parent_succ)
            else:
                result = _delete_leaf(succ, parent_succ)
            
            ### END SOLUTION
            
            return result
        
        # delete main function ...
        # -----------------------------------------
        
        # keep track of the result of deletion ...
        result = False
        
        # find it ...
        node, parent = BSTNode.find(self, None, key)
        
        if node is not None:
            # found it ... what type of BSTNode?
            nrChildren = node.child_count()
                    
            if nrChildren == 0:
                # ... it's a leaf BSTNode ...
                result = _delete_leaf(node, parent)
            elif nrChildren == 1:
                # it has a single child ...
                result = _delete_one(node, parent)
            else: # nrChildren == 2
                result = _delete_two(node, parent)
            
        # give back if succeeded or not ...
        return result               


    def traverse_tree_pre_order(self, visitor):
        """traverese tree pre-order mode and fire visitor on BSTNode"""
        
        ### BEGIN SOLUTION
        
        # ... first, process the BSTNode ...
        visitor(self)
       
        # ... then left tree ...
        if self.left is not None:
            self.left.traverse_tree_pre_order(visitor)
            
        # ... finally, right tree...
        if self.right is not None:       
            self.right.traverse_tree_pre_order(visitor)
        
        ### END SOLUTION
            
    def traverse_tree_in_order(self, visitor):
        """traverese tree in-order mode and fire visitor on BSTNode"""
        
        # left tree first ...
        if self.left is not None:
            self.left.traverse_tree_in_order(visitor)
            
        # ... process the BSTNode in-order ...
        visitor(self)
        
        # ... right tree...
        if self.right is not None:       
            self.right.traverse_tree_in_order(visitor)


    def traverse_tree_post_order(self, visitor):
        """traverese tree post-order mode and fire visitor on BSTNode"""
        
        ### BEGIN SOLUTION
        
        # ... first, left tree ...
        if self.left is not None:
            self.left.traverse_tree_post_order(visitor)
            
        # ... then, right tree...
        if self.right is not None:       
            self.right.traverse_tree_post_order(visitor)
            
        # ... finally, process the BSTNode ...
        visitor(self)
        
        ### END SOLUTION
        
        
```

<!-- grade_id: at_1_0_1, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true

ONE = 1
assert_true( ONE == 1 )

### END HIDDEN TESTS
```

#### Task 1-1

Complete the implementation of ` def insert(self, data)` and run the below function `test_add_names()` to verify your insert. Check both the names, and their position in the tree, given the order of insertion. The above given diagram is showing the tree you should obtain.

<!-- grade_id: q_1_1_0, points: 0, grade: false, solution: true -->
```python
# Task 1.1 code cell

from BSTPrinter import *

# Define some variables and methods for testing...
ROOT  = 'Lars'  
NAMES = [
         'Maze', 'Leon', 'John', 'Zack', 'Chuc', 'Pete', 'Lisa', \
         'Mary', 'Rodd', 'Bert', 'Sean', 'Will', 'Dora', 'Anna'  
        ]
ALL_NAMES = NAMES + [ ROOT ]

# factory function making a BST ...
def make_bst():
    # add the root name ...
    a_bst = BSTNode( ROOT )
    
    # add all other names in addition to the root ...
    
    for name in NAMES:
        ### BEGIN SOLUTION
        a_bst.insert(name)
        ### END SOLUTION
    
    return a_bst


# Tests are down here...
def test_add_names():
    
    BSTPrinter.get_instance().print_header('Test whether all names correctly added to bst')
    bst = make_bst()

    # print a list of nodes in-order ...
    BSTPrinter.get_instance().print_BSTNodes_sorted(bst, 'list of tree nodes, in-order', True)
    
    # print the structure of the tree ...
    BSTPrinter.get_instance().print_tree_structure(bst, 'bst tree structure after insertion names')

    
## run the test ...
test_add_names()

```

Output:
```
--------------------------------------------------------------------------------
Test whether all names correctly added to bst
--------------------------------------------------------------------------------
----------------------------------------
list of tree nodes, in-order
----------------------------------------
Anna Bert Chuc Dora John Lars Leon Lisa Mary Maze Pete Rodd Sean Will Zack 
----------------------------------------
bst tree structure after insertion names
----------------------------------------
            Zack
                                    Will
                              Sean
                        Rodd
                  Pete
      Maze
                        Mary
                  Lisa
            Leon
Lars
      John
                  Dora
            Chuc
                  Bert
                        Anna
```

<!-- grade_id: at_1_1_0, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true

test_bst = make_bst()

assert_true( test_bst.data == ROOT )
assert_true( test_bst.left.data == 'John')
assert_true( test_bst.left.left.data == 'Chuc')
assert_true( test_bst.left.left.right.data == 'Dora')
assert_true( test_bst.left.left.left.left.data == 'Anna')

### END HIDDEN TESTS
```

<!-- grade_id: at_1_1_1, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true

test_bst = make_bst()

assert_true( test_bst.right.data == 'Maze' )
assert_true( test_bst.right.left.data == 'Leon' )
assert_true( test_bst.right.right.data == 'Zack' )
assert_true( test_bst.right.right.right is None )
assert_true( test_bst.right.right.left.right.right.right.data == 'Will' )
assert_true( test_bst.right.right.left.right.right.left is None )

### END HIDDEN TESTS
```

#### Task 1.2

Finalize static function `BSTNode.find(bst, None, name)` and find back all names inserted using the below function.

<!-- grade_id: q_1_2_0, points: 0, grade: false, solution: true -->
```python
# Task 1.2 code cell

def find_all_names(bst, namelist):
    if bst is None:
        return
    for name in namelist:
        
        # find name in the bst and print the result ...
        
        ### BEGIN SOLUTION
        
        node, parent = BSTNode.find(bst, None, name) 
        
        ### END SOLUTION
        
        if node is None:
            print('{:s} not on tree'. format(str(name)))
        else:
            print('{:s} found: {:s}'. \
                  format(str(name), str(node)), end='')
            if parent is not None:
                print(' (child of: {:s})'. format(str(parent)) )
            else:
                print(' (no parent)')


# Tests are down here...
def test_find_names():
    BSTPrinter.get_instance().print_header('Test whether all names inserted can be found on bst')
    
    bst = make_bst()

    # Create some names that are not in the tree
    NOT_ON_TREE = ['Jacques', 'Kit']
    TEST_LIST = ALL_NAMES + NOT_ON_TREE

    # For all names, display whether they are in the tree
    find_all_names(bst, TEST_LIST)

    
## run the test ...
test_find_names()

```

Output:
```
--------------------------------------------------------------------------------
Test whether all names inserted can be found on bst
--------------------------------------------------------------------------------
Maze found: Maze (child of: Lars)
Leon found: Leon (child of: Maze)
John found: John (child of: Lars)
Zack found: Zack (child of: Maze)
Chuc found: Chuc (child of: John)
Pete found: Pete (child of: Zack)
Lisa found: Lisa (child of: Leon)
Mary found: Mary (child of: Lisa)
Rodd found: Rodd (child of: Pete)
Bert found: Bert (child of: Chuc)
Sean found: Sean (child of: Rodd)
Will found: Will (child of: Sean)
Dora found: Dora (child of: Chuc)
Anna found: Anna (child of: Bert)
Lars found: Lars (no parent)
Jacques not on tree
Kit not on tree
```

<!-- grade_id: at_1_2_0, points: 4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true

# check the root node ...
test_bst = make_bst()

test_names = ('Lars', 'Maze', 'Leon', 'John', 'Zack', 'Chuc', 'Pete', 'Lisa',
         'Mary', 'Rodd', 'Bert', 'Sean', 'Will', 'Dora', 'Anna' )

not_on_tree_names = ('Kitt', 'Tovi', 'Burt')

child_parent = (
    ('Maze', 'Lars'),
    ('Leon', 'Maze'),
    ('John', 'Lars'),
    ('Zack', 'Maze'),
    ('Chuc', 'John'),
    ('Pete', 'Zack'),
    ('Lisa', 'Leon'),
    ('Mary', 'Lisa'),
    ('Rodd', 'Pete'),
    ('Bert', 'Chuc'),
    ('Sean', 'Rodd'),
    ('Will', 'Sean'),
    ('Dora', 'Chuc'),
    ('Anna', 'Bert')
)

for name in test_names + not_on_tree_names:
    node, parent = BSTNode.find(test_bst, None, name)
    if node is None:
        assert_true( name in not_on_tree_names and name not in test_names )
    elif parent is None:
        assert_true( name == 'Lars' )
    else:
        assert_true( (node.data, parent.data) in child_parent )

### END HIDDEN TESTS
```

#### Task 1.3

Finalize the implementation of static function `height()` in `class BSTNode`. Run the test routine `test_height()` below, to test it.

<!-- grade_id: q_1_3_0, points: 0, grade: false, solution: true -->
```python
# Task 1.3 code cell

def test_height():
    BSTPrinter.get_instance().print_header('Test height computation of bst')
    bst = make_bst()
    
    ### BEGIN SOLUTION
    
    bst_height = BSTNode.height(bst)
    
    ### END SOLUTION
    
    print("Height must be 7, actual height: {:d}".format(bst_height))
    
    
#run the test ...
test_height()
```

Output:
```
--------------------------------------------------------------------------------
Test height computation of bst
--------------------------------------------------------------------------------
Height must be 7, actual height: 7
```

<!-- grade_id: at_1_3_0, points: 4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

from nose.tools import assert_true

test_bst = make_bst()

assert_true( BSTNode.height(test_bst) == 7 )

### END HIDDEN TESTS
```

#### Task 1.4

Delete nodes with no, one child or two children. We do this in two steps:

* step 1: implement `_delete_leaf()` and `_delete_one()`, having zero and one child, resp. Test these first
* step 2: implement `_delete_two()`, which is the most difficult one

**Hints for this Exercise**

* for an explanation of what `delete` does, consult the lecture slides
* when studying the templates given, you will notice that we opted to find and promote a *successor* in the deletion of a 2-children-node (and not a predecessor). A predecessor-choice would have been proceeding along the same lines, however
* further notice that after swapping tree nodes (successor --> node-to-delete), the *successor node* is itself being deleted by a call to `_delete_one(successor, parent_successor)`. You can argue that deleting a *successor node* can be atmost a 1-child-deletion
* observe that the hierarchy of puter-and-inner functions is as follllows:

```
    def delete(self, key)
        """delete node with key given"""
        
        # aux functions for delete ...
        def _delete_leaf(node, parent):
            """delete a leaf node; 0 children"""
                 
        def _delete_one(node, parent):
            """delete a node with 1 child"""

        def _delete_two(node, parent):
            """delete a node with two subtrees"""
            
            def _find_successor(node, parent):
                """find successor and keep track of node and parent"""
```

* all the above functions return `True` if the node was found and deleted succesfully, and `False` otherwise, except function `_find_successor()`, which return the successor found and its parent.


Step 1: delete nodes with no or one child. Finalize `_delete_leaf(node, parent)` and `_delete_one(node, parent)` in function `delete(self, key)`. Leave `_delete_two(node, parent)` arest for now. Run the below code and test the results. Check the printed tree diagram. Experiment with names you select, below.

<!-- grade_id: q_1_4_0, points: 0, grade: false, solution: true -->
```python
# Task 1.4 code cell

def test_delete_names_01():
    BSTPrinter.get_instance().print_header('Test deleting nodes with one or zero children')
    bst = make_bst()

    # delete a few ...
    for name in ('Sean', 'Anna', 'John', 'Lisa'):
        # delete from the tree ...
        
        ### BEGIN SOLUTION
        
        bst.delete(name)
        
        ### END SOLUTION

    # check the result ...
    find_all_names(bst, ALL_NAMES)

    # print all the nodes in order ...
    BSTPrinter.get_instance().print_tree_structure(bst, 'bst after deletions one and zero children')

# run the test  ...
test_delete_names_01()
```

Output:
```
--------------------------------------------------------------------------------
Test deleting nodes with one or zero children
--------------------------------------------------------------------------------
Maze found: Maze (child of: Lars)
Leon found: Leon (child of: Maze)
John not on tree
Zack found: Zack (child of: Maze)
Chuc found: Chuc (child of: Lars)
Pete found: Pete (child of: Zack)
Lisa not on tree
Mary found: Mary (child of: Leon)
Rodd found: Rodd (child of: Pete)
Bert found: Bert (child of: Chuc)
Sean not on tree
Will found: Will (child of: Rodd)
Dora found: Dora (child of: Chuc)
Anna not on tree
Lars found: Lars (no parent)
----------------------------------------
bst after deletions one and zero children
----------------------------------------
            Zack
                              Will
                        Rodd
                  Pete
      Maze
                  Mary
            Leon
Lars
            Dora
      Chuc
            Bert
```

<!-- grade_id: at_1_4_0, points: 4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true

def test_find(root, parent, key):
    """return node with key given and its parent, or None if not found"""
    # key in the root?
    if root is None or root.data == key:
        return root, parent
    # continue searching down the tree ...
    # key greater than root's key?
    if key > root.data:
        return test_find(root.right, root, key)
    # Key is smaller than root's key
    return test_find(root.left, root, key)

def test_height(tree):
    tree_height = 0
    if tree is not None:
        height_left  = test_height(tree.left)
        height_right = test_height(tree.right)
        tree_height  =max(height_left, height_right) + 1 
    return tree_height
    
test_bst = make_bst()

height = 7

for name in ('Sean', 'Anna', 'John', 'Lisa'):
    test_bst.delete(name)
    # try to find it back ...
    node, parent = test_find(test_bst, None, name)
    assert_true( node is None )

# do extra deletion ...
new_height = test_height(test_bst)
assert_true( new_height <= height )
height = new_height

node, parent = test_find(test_bst, None, 'Will')
if node is not None:
    test_bst.delete('Will')
    new_height = test_height(test_bst)
    assert_true( new_height == height-1 )

### END HIDDEN TESTS
```

#### Task 1.5

Step 2: delete nodes with two children. Finalize `_delete_two(node, parent)` and `_find_successor(node, parent)` in function `delete(self, key)`. Complete the call to this method for nodes with two children. Run the below code and test the results. Check the printed tree diagram. Experiment with names you select, below.

<!-- grade_id: q_1_5_0, points: 0, grade: false, solution: true -->
```python
# Task 1.5 code cell

def test_delete_names_02():
    BSTPrinter.get_instance().print_header('Test deleting nodes with two children')
    bst = make_bst()

    # delete a few ...
    for name in ('Lars', 'Maze', 'Chuc'):
        # delete from the tree ...
        
        ### BEGIN SOLUTION
        
        bst.delete(name)
        
        ### END SOLUTION

    # check the result ...
    find_all_names(bst, ALL_NAMES)

    # print all the nodes in order ...
    BSTPrinter.get_instance().print_tree_structure(bst, 'bst after deletions two children nodes')


# run the test ...
test_delete_names_02()
```

Output:
```
--------------------------------------------------------------------------------
Test deleting nodes with two children
--------------------------------------------------------------------------------
Maze not on tree
Leon found: Leon (no parent)
John found: John (child of: Leon)
Zack found: Zack (child of: Pete)
Chuc not on tree
Pete found: Pete (child of: Leon)
Lisa found: Lisa (child of: Pete)
Mary found: Mary (child of: Lisa)
Rodd found: Rodd (child of: Zack)
Bert found: Bert (child of: Dora)
Sean found: Sean (child of: Rodd)
Will found: Will (child of: Sean)
Dora found: Dora (child of: John)
Anna found: Anna (child of: Bert)
Lars not on tree
----------------------------------------
bst after deletions two children nodes
----------------------------------------
            Zack
                              Will
                        Sean
                  Rodd
      Pete
                  Mary
            Lisa
Leon
      John
            Dora
                  Bert
              
```

<!-- grade_id: at_1_5_0, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true

test_bst = make_bst()
height = test_height(test_bst)

# delete the root ...
name = ROOT
test_bst.delete(name)
# try to find it back ...
node, parent = test_find(test_bst, None, name)
assert_true( node is None )
new_height = test_height(test_bst)
assert_true( new_height <= height )

### END HIDDEN TESTS
```

<!-- grade_id: at_1_5_1, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true

test_bst = make_bst()
height = test_height(test_bst)

# delete the root ...
name = ROOT
test_bst.delete(name)

# find the new root
new_root = test_bst.data
# look it up and see parent is None
new_root_node, new_root_parent = test_find(test_bst, None, new_root)
assert_true( new_root_node is not None )
assert_true( new_root_node.data == new_root )
assert_true( new_root_parent is None )

### END HIDDEN TESTS
```

<!-- grade_id: at_1_5_2, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true

test_name = 'Pamm'
test_bst = make_bst()
height = test_height(test_bst)


# add a successor to the right subtree ... 
test_bst.insert(test_name) # the new succssor ...

# BSTPrinter.get_instance().print_tree_structure(test_bst,'test tree with test_name')

# delete Maze (who has 2 children: Leon and Zack. Successor: test_name)
test_bst.delete('Maze')

# BSTPrinter.get_instance().print_tree_structure(test_bst,'test tree with test_name, without Maze')

right_of_root_node, root_node = test_find(test_bst, None, test_name)
assert_true(root_node.right is right_of_root_node)
assert_true(right_of_root_node.data == test_name)
assert_true(right_of_root_node.right.data == 'Zack')
assert_true(right_of_root_node.left.data == 'Leon')

### END HIDDEN TESTS
```

#### Task 1.6

Implement `traverse_tree_pre_order()` and `traverse_tree_post_order()` and run it below, with `BSTNode.print_BSTNode` as the visitor function. Check if the pre- and post-order produce the correct sequence of nodes.

<!-- grade_id: q_1_6_0, points: 0, grade: false, solution: true -->
```python
# Task 1.6 code cell

def test_traversers():
    BSTPrinter.get_instance().print_header('Test if tree traversed correctly, pre- and post-order')
    bst = make_bst()

    # check the result ...
    print('tree nodes in pre-order traversal with visitor: BSTNode.print_BSTNode')
    
    ### BEGIN SOLUTION
    
    bst.traverse_tree_pre_order(BSTNode.print_BSTNode); print()
    
    ### END SOLUTION
    
    print('tree nodes in post-order traversal with visitor: BSTNode.print_BSTNode')
    
    ### BEGIN SOLUTION
    
    bst.traverse_tree_post_order(BSTNode.print_BSTNode); print()
    
    ### END SOLUTION

# run the test ...
test_traversers()
```

Output:
```
--------------------------------------------------------------------------------
Test if tree traversed correctly, pre- and post-order
--------------------------------------------------------------------------------
tree nodes in pre-order traversal with visitor: BSTNode.print_BSTNode
Lars John Chuc Bert Anna Dora Maze Leon Lisa Mary Zack Pete Rodd Sean Will 
tree nodes in post-order traversal with visitor: BSTNode.print_BSTNode
Anna Bert Dora Chuc John Mary Lisa Leon Will Sean Rodd Pete Zack Maze Lars
```

<!-- grade_id: at_1_6_0, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true

# mock output ...
assert_true( 1==1 )

### END HIDDEN TESTS
```

# EXERCISE 2: implement a `Heap`
A heap is an array-based data structure in which items are organized using a hierarchical model with a min-or-max parent-child relationship. This invariant relation is maintained while pushing new items on the heap, or popping items from the heap. This again implies that we can easily obtain the next-in-order item from a heap, no matter how its content changes. This makes a heap particularly apt for sorting, but also for dynamic applications, such as a priority-queue.

<!-- grade_id: q_2_0_0, points: 0, grade: false, solution: true -->
```python
class Heap(list):
    """zero-indexed list-based min-or-max heap class"""
    
    # distinguish between ordering types ...
    MIN_HEAP, MAX_HEAP = 0, 1

    @staticmethod
    def print_heap(a_heap, title=''):
        '''print heap as a string'''
        print('{:s} {:s}'.format(title, str(a_heap)))


    def __init__(self, hp_type=MIN_HEAP, items=None):
        '''init a heap as a list'''
        super().__init__(self)
        # register the heap ordering ...
        self.type = hp_type
        # if items given, preload on heap ...
        if items:
            self.extend(items)
        # install heap on list ...
        self.heapify()
        
    def __str__(self):
        """return list as string"""
        _items = ', '.join([str(_item) for _item in self])
        return _items

    def empty(self):
        """return T|F heap empty"""
        
        ### BEGIN SOLUTION
        
        return not self.size() > 0
    
        ### END SOLUTION
    
    def size(self):
        '''return current heap size'''
        
        ### BEGIN SOLUTION
        
        return len(self)
    
        ### END SOLUTION

    def push_heap(self, item):
        ''' push new item on the heap'''
        # append the item ...
        
        ### BEGIN SOLUTION
        
        self.append(item)
        
        ### END SOLUTION
        
        # bubble up ...
        
        ### BEGIN SOLUTION
        
        self._bubble_up(self.size()-1)
        
        ### END SOLUTION

    def pop_heap(self):
        '''pop from the top of the heap'''
        # pop heap top [0] and swap with last ...
        
        ### BEGIN SOLUTION
        
        item, self[0] = self[0], self[-1]
        
        ### END SOLUTION
        
        # remove the double entry ...
        
        ### BEGIN SOLUTION
        
        dupl = self.pop(self.size()-1)
        
        ### END SOLUTION
        
        # bubble down ...
        
        ### BEGIN SOLUTION
        
        self._bubble_down(0)
        
        ### END SOLUTION
        
        return item

    def in_order(self, child, parent):
        '''return T|F parent and child are in heap-order'''
        if parent is None or child is None:
            return True
        
        ### BEGIN SOLUTION
        
        if self.type == Heap.MAX_HEAP:
            return self[parent] >= self[child]
        else:
            return self[parent] <= self[child]
        
        ### END SOLUTION

    def _bubble_up(self, child):
        '''swap last inserted element into correct position'''
        # swap child-parent until parent in order ...
        
        ### BEGIN SOLUTION
        
        parent = (child-1) // 2
        while parent >= 0 and not self.in_order(child, parent):
            # swap them ...
            self[child], self[parent] = self[parent], self[child]
            child = parent
            parent = (child - 1) // 2
        
        ### END SOLUTION
        

    def _bubble_down(self, parent):
        '''swap parent downwards into correct position'''
        def child_to_swap(lchild, rchild, size):
            '''return index of child to involve in swapping, None if no child'''
            chld = None
            
            ### BEGIN SOLUTION
            
            if lchild < size:
                chld = lchild
            if rchild < size:
                if self.type == Heap.MAX_HEAP:
                    if self[rchild] > self[lchild]:
                        chld = rchild
                else:
                    # it's a MIN_HEAP ...
                    if self[rchild] < self[lchild]:
                        chld = rchild
            
            ### END SOLUTION
            
            return chld

        if self.empty():
            return
        size   = self.size()
        lchild  = 2*parent + 1
        rchild = lchild + 1
        swp_ndx= child_to_swap(lchild, rchild, size)

        ### BEGIN SOLUTION
        
        while not self.in_order(swp_ndx, parent):
            # bubble down parent, swap child upward ...
            # the swapped child is always the greatest ...
            self[parent], self[swp_ndx] = self[swp_ndx], self[parent]
            parent = swp_ndx
            lchild = 2*parent + 1
            rchild = lchild + 1
            swp_ndx= child_to_swap(lchild, rchild, size)
        
        ### END SOLUTION

    def heapify(self):
        '''install heap property on list'''
        size = self.size()
        # heap tree leaves have no children and 
        # already have heap property ...
        # skip leaves ...
        for ndx in range (size // 2, -1, -1):
            self._bubble_down(ndx)
        
```

<!-- grade_id: at_2_0_0, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true

# mock output ...
assert_true( 1==1 )

### END HIDDEN TESTS
```

#### Task 2.1
Complete function member `Heap.empty()` and `Heap.size()`, plus the below test code. Run it and verify the results.

**Note**: until after Task 2.2, the ordering in the Heap may NOT be OK, but the count should.

<!-- grade_id: q_2_1_0, points: 0, grade: false, solution: true -->
```python
# Task 2.1 code cell

print('test creation of min or max heaps, from item lists, different item types')

def test_make_heap():
    """create suite of heaps, and print"""
    none_list  = None
    empty_list = []
    int_list   = [10,7,2,3,9,4,8,1,0,5,6]
    char_list  = ['p','a','h','k','q','r','l','f','o',
                    'g','s','d','m','e','v','t','b','j','w',
                                  'z','x','y','i','u','c','n']
    dupl_list  = [5.0, 2.1, 5.0, 2.2, 7.6, 4.2, 2.1, 5.0, 4.3]
    
    # create heaps ...
    none_heap      = Heap(hp_type=Heap.MAX_HEAP, items=none_list)
    empty_heap     = Heap(hp_type=Heap.MIN_HEAP, items=empty_list)
    # make int_max_heap and char_min_heap heap ...
    
    ### BEGIN SOLUTION
    
    int_max_heap   = Heap(hp_type=Heap.MAX_HEAP, items=int_list)
    char_min_heap  = Heap(hp_type=Heap.MIN_HEAP, items=char_list)
    
    ### END SOLUTION
    
    dupl_heap      = Heap(items=dupl_list) # default is: min heap
    default_heap   = Heap()
    
    for k, hp in enumerate( (none_heap, empty_heap, int_max_heap, char_min_heap, dupl_heap, default_heap) ):
        print(f'heap: {k:d}, size={hp.size():2d}: [{str(hp):s}] Heap empty? {str(hp.empty())}')

test_make_heap()
```

Output:
```
test creation of min or max heaps, from item lists, different item types
heap: 0, size= 0: [] Heap empty? True
heap: 1, size= 0: [] Heap empty? True
heap: 2, size=11: [10, 9, 8, 3, 7, 4, 2, 1, 0, 5, 6] Heap empty? False
heap: 3, size=26: [a, b, c, f, g, d, e, k, j, q, i, h, m, l, v, t, p, o, w, z, x, y, s, u, r, n] Heap empty? False
heap: 4, size= 9: [2.1, 2.2, 2.1, 4.3, 7.6, 4.2, 5.0, 5.0, 5.0] Heap empty? False
heap: 5, size= 0: [] Heap empty? True
```

<!-- grade_id: at_2_1_0, points: 3, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true

# test_lists
test_none_list  = None
test_empty_list = []
test_int_list   = [10,7,2,3,9,4,8,1,0,5,6]
test_char_list  = ['p','a','h','k','q','r','l','f','o',
                    'g','s','d','m','e','v','t','b','j','w',
                                  'z','x','y','i','u','c','n']
test_dupl_list  = [5.0, 2.1, 5.0, 2.2, 7.6, 4.2, 2.1, 5.0, 4.3]
    
# create heaps ...
test_none_heap      = Heap(hp_type=Heap.MAX_HEAP, items=test_none_list)
test_empty_heap     = Heap(hp_type=Heap.MIN_HEAP, items=test_empty_list)
test_int_max_heap   = Heap(hp_type=Heap.MAX_HEAP, items=test_int_list)
test_char_min_heap  = Heap(hp_type=Heap.MIN_HEAP, items=test_char_list)
test_dupl_heap      = Heap(items=test_dupl_list) # default is: min heap
test_default_heap   = Heap()

assert_true( test_none_heap.size()    ==  0 )
assert_true( test_empty_heap.size()   ==  0 )
assert_true( test_int_max_heap.size() == 11 )
assert_true( test_char_min_heap.size()== 26 )
assert_true( test_dupl_heap.size()    ==  9 )
assert_true( test_default_heap.size() ==  0 )
        
### END HIDDEN TESTS
```

<!-- grade_id: at_2_1_1, points: 3, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true

# test_lists
test_none_list  = None
test_empty_list = []
test_int_list   = [10,7,2,3,9,4,8,1,0,5,6]
test_char_list  = ['p','a','h','k','q','r','l','f','o',
                    'g','s','d','m','e','v','t','b','j','w',
                                  'z','x','y','i','u','c','n']
test_dupl_list  = [5.0, 2.1, 5.0, 2.2, 7.6, 4.2, 2.1, 5.0, 4.3]
    
# create heaps ...
test_none_heap      = Heap(hp_type=Heap.MAX_HEAP, items=test_none_list)
test_empty_heap     = Heap(hp_type=Heap.MIN_HEAP, items=test_empty_list)
test_int_max_heap   = Heap(hp_type=Heap.MAX_HEAP, items=test_int_list)
test_char_min_heap  = Heap(hp_type=Heap.MIN_HEAP, items=test_char_list)
test_dupl_heap      = Heap(items=test_dupl_list) # default is: min heap
test_default_heap   = Heap()

assert_true( test_none_heap.empty()    == True  )
assert_true( test_empty_heap.empty()   == True  )
assert_true( test_int_max_heap.empty() == False )
assert_true( test_char_min_heap.empty()== False )
assert_true( test_dupl_heap.empty()    == False )
assert_true( test_default_heap.empty() == True  )

### END HIDDEN TESTS
```

#### Task 2.2
Implement `Heap.in_order(self, child, parent)` and run the below code to test

<!-- grade_id: q_2_2_0, points: 0, grade: false, solution: true -->
```python
# Task 2.2 code cell

print('test order in heaps to be correct')
def test_in_order(heap):
    """test if items in heap obey order"""
    #  walk up the heap ...
    for ndx in range (heap.size() // 2, -1, -1):
        parent = (ndx-1) // 2
        if parent >= 0:
            if heap.type == Heap.MIN_HEAP:
                if heap[parent] <= heap[ndx] and heap[parent] <= heap[ndx+1]:
                    print(f'min heap: parent {heap[parent]} <= {heap[ndx]}, {heap[ndx+1]}')
                else:
                    print(f'min heap: ordering error: parent {heap[parent]}, children: {heap[ndx]}, {heap[ndx+1]}')
            else:
                ### BEGIN SOLUTION
                
                if heap[parent] >= heap[ndx] and heap[parent] >= heap[ndx+1]:
                    print(f'max heap: parent {heap[parent]} >= {heap[ndx]}, {heap[ndx+1]}')
                else:
                    print(f'max heap: ordering error: parent {heap[parent]}, children: {heap[ndx]}, {heap[ndx+1]}')
                
                ### END SOLUTION

int_min_heap = Heap(hp_type=Heap.MIN_HEAP, items=[9, 3, 5, 8, 2, 1, 0, 4, 7, 6])
chr_max_heap = Heap(hp_type=Heap.MAX_HEAP, items=['f','u','r','d','s','f','f','h'])

for hp in (int_min_heap, chr_max_heap):
    test_in_order(hp)
```

Output:
```
test order in heaps to be correct
min heap: parent 1 <= 9, 5
min heap: parent 2 <= 3, 9
min heap: parent 2 <= 4, 3
min heap: parent 0 <= 1, 4
min heap: parent 0 <= 2, 1
max heap: parent s >= f, f
max heap: parent s >= h, f
max heap: parent u >= r, h
max heap: parent u >= s, r
```

<!-- grade_id: at_2_2_0, points: 3, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true, assert_false

def test_in_order(htype, child, parent):
    '''return T|F parent and child are in heap-order'''
    if parent is None or child is None:
        return True
    if htype == Heap.MAX_HEAP:
        return parent_val >= child_val
    else:
        return parent_val <= child_val

my_min_heap = Heap(Heap.MIN_HEAP, [2.1, 2.1, 2.2, 4.2, 4.3, 5.0, 5.0, 5.0, 7.6])
my_max_heap = Heap(Heap.MAX_HEAP, [10, 7, 2, 3, 9, 4, 8, 1, 0, 5, 6])

for hp in ( my_min_heap, my_max_heap):
    # Heap.print_heap(hp)
    htype = hp.type
    size  = hp.size()
    size = 0 # failure
    for parent in range( 0, (size-1) // 2):
        lchild = 2*parent + 1
        rchild = lchild   + 1
        if lchild < size-1 and rchild < size-1:
            lchild_val = hp[lchild]
            rchild_val = hp[rchild]
            parent_val = hp[parent] 
            # print(f'{htype}, index: parent: {parent}, child: {lchild} and {rchild}')
            # print(f'{htype}, value: parent: {parent_val}, child: {lchild_val} and {rchild_val}')
            assert_true( hp.in_order(lchild, parent) == test_in_order(htype, lchild_val, parent_val)) 
            assert_true( hp.in_order(rchild, parent) == test_in_order(htype, rchild_val, parent_val)) 
                    
### END HIDDEN TESTS
```

#### Task 2.3
Implement `Heap._bubble_down()`, bubbling down parents and children to restore the heap ordering. To test your implementation, change in `class Heap` above:

```python
    def heapify(self):
        '''install heap property on list'''
        size = self.size()
        # heap tree leaves have no children and 
        # already have heap property ...
        # skip leaves ...
        for ndx in range (size // 2, -1, -1):
            #self._bubble_down(ndx)
            pass
```

to:

```python
    def heapify(self):
        '''install heap property on list'''
        size = self.size()
        # heap tree leaves have no children and 
        # already have heap property ...
        # skip leaves ...
        for ndx in range (size // 2, -1, -1):
            self._bubble_down(ndx)
```

first **RERUN the class code above**, then run the below test code. It should produce heaps with the correct ordering now. Verify this using the below code.

<!-- grade_id: q_2_3_0, points: 0, grade: false, solution: true -->
```python
# Task 2.3 code cell

print('test creation of min or max heaps,  different item types, with the correct heap ordering')

def test_create_heap_in_order():
    """create suite of heaps, and print"""
    none_list  = None
    empty_list = []
    int_list   = [10,7,2,3,9,4,8,1,0,5,6]
    char_list  = ['p','a','h','k','q','r','l','f','o',
                    'g','s','d','m','e','v','t','b','j','w',
                                  'z','x','y','i','u','c','n']
    dupl_list  = [5.0, 2.1, 5.0, 2.2, 7.6, 4.2, 2.1, 5.0, 4.3]
    
    # create heaps ...
    ### BEGIN SOLUTION
    
    none_heap      = Heap(hp_type=Heap.MAX_HEAP, items=none_list)
    empty_heap     = Heap(hp_type=Heap.MIN_HEAP, items=empty_list)
    int_max_heap   = Heap(hp_type=Heap.MAX_HEAP, items=int_list)
    char_min_heap  = Heap(hp_type=Heap.MIN_HEAP, items=char_list)
    dupl_heap      = Heap(items=dupl_list) # default is: min heap
    default_heap   = Heap()
    
    ### END SOLUTION
    
    for k, hp in enumerate( (none_heap, empty_heap, int_max_heap, char_min_heap, dupl_heap, default_heap) ):
        print(f'heap: {k:d}, size={hp.size():2d}: [{str(hp):s}]')
    
test_create_heap_in_order()
```

Output:
```
test creation of min or max heaps,  different item types, with the correct heap ordering
heap: 0, size= 0: []
heap: 1, size= 0: []
heap: 2, size=11: [10, 9, 8, 3, 7, 4, 2, 1, 0, 5, 6]
heap: 3, size=26: [a, b, c, f, g, d, e, k, j, q, i, h, m, l, v, t, p, o, w, z, x, y, s, u, r, n]
heap: 4, size= 9: [2.1, 2.2, 2.1, 4.3, 7.6, 4.2, 5.0, 5.0, 5.0]
heap: 5, size= 0: []
```

<!-- grade_id: at_2_3_0, points: 4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true, assert_false

# how to test this ...
assert_true( 1==1 )
                    
### END HIDDEN TESTS
```

#### Task 2.4
Implement `Heap.pop_heap()` using `Heap.bubble_down()`. Run the below code to verify the result.

<!-- grade_id: q_2_4_0, points: 0, grade: false, solution: true -->
```python
# Task 2.4 code cell

print('test popping of min or max heaps,  different item types, with the correct heap ordering')

def test_pop_heap():
    """create suite of heaps, and print"""
    none_list  = None
    empty_list = []
    int_list   = [10,7,2,3,9,4,8,1,0,5,6]
    char_list  = ['p','a','h','k','q','r','l','f','o',
                    'g','s','d','m','e','v','t','b','j','w',
                                  'z','x','y','i','u','c','n']
    dupl_list  = [5.0, 2.1, 5.0, 2.2, 7.6, 4.2, 2.1, 5.0, 4.3]
    
    # create heaps ...
    none_heap      = Heap(hp_type=Heap.MAX_HEAP, items=none_list)   # heap 0
    empty_heap     = Heap(hp_type=Heap.MIN_HEAP, items=empty_list)  # heap 1
    int_max_heap   = Heap(hp_type=Heap.MAX_HEAP, items=int_list)    # ...
    char_min_heap  = Heap(hp_type=Heap.MIN_HEAP, items=char_list)
    dupl_heap      = Heap(items=dupl_list) # default is: min heap
    
    for k, hp in enumerate( (none_heap, empty_heap, int_max_heap, char_min_heap, dupl_heap) ):
        print(f'heap {k:d} popped: ', end='')
        while not hp.empty():
            # pop the next item amd print it ...
            
            ### BEGIN SOLUTION
            
            item = hp.pop_heap()
            
            ### END SOLUTION
            
            print(f'{item}', end=' ')
        print('')

test_pop_heap()
```

Output:
```
test popping of min or max heaps,  different item types, with the correct heap ordering
heap 0 popped: 
heap 1 popped: 
heap 2 popped: 10 9 8 7 6 5 4 3 2 1 0 
heap 3 popped: a b c d e f g h i j k l m n o p q r s t u v w x y z 
heap 4 popped: 2.1 2.1 2.2 4.2 4.3 5.0 5.0 5.0 7.6
```

<!-- grade_id: at_2_4_0, points: 4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true, assert_false

def test_pop_min_heap_correct():
    """create suite of heaps, and print"""
    test_int_list   = [10,7,2,3,9,4,8,1,0,5,6]
    test_char_list  = ['p','a','h','k','q','r','l','f','o',
                    'g','s','d','m','e','v','t','b','j','w',
                                  'z','x','y','i','u','c','n']
    test_dupl_list  = [5.0, 2.1, 5.0, 2.2, 7.6, 4.2, 2.1, 5.0, 4.3]
    
    # create heaps ...
    test_int_max_heap   = Heap(hp_type=Heap.MAX_HEAP, items=test_int_list)    # ...
    test_char_min_heap  = Heap(hp_type=Heap.MIN_HEAP, items=test_char_list)
    test_dupl_heap      = Heap(items=test_dupl_list) # default is: min heap
    
    # test MIN HEAP types first...
    for hp in (test_char_min_heap, test_dupl_heap):
        my_max = hp.pop_heap()
        while not hp.empty():
            item = hp.pop_heap()
            assert_true( item >= my_max )
            my_max = item
        
test_pop_min_heap_correct()

### END HIDDEN TESTS
```

<!-- grade_id: at_2_4_1, points: 4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true, assert_false

def test_pop_max_heap_correct():
    """create suite of heaps, and print"""
    test_int_list   = [10,7,2,3,9,4,8,1,0,5,6]
    test_char_list  = ['p','a','h','k','q','r','l','f','o',
                    'g','s','d','m','e','v','t','b','j','w',
                                  'z','x','y','i','u','c','n']
    test_dupl_list  = [5.0, 2.1, 5.0, 2.2, 7.6, 4.2, 2.1, 5.0, 4.3]
    
    # create heaps ...
    test_int_max_heap   = Heap(hp_type=Heap.MAX_HEAP, items=test_int_list)    # ...
    test_char_min_heap  = Heap(hp_type=Heap.MIN_HEAP, items=test_char_list)
    test_dupl_heap      = Heap(items=test_dupl_list) # default is: min heap
    
    # test MAX HEAP types first...
    for hp in (test_int_max_heap, ):
        my_min = hp.pop_heap()
        while not hp.empty():
            item = hp.pop_heap()
            assert_true( item <= my_min )
            my_min = item
        
test_pop_max_heap_correct()

### END HIDDEN TESTS
```

#### Task 2.5
Implement `Heap.push_heap()` and `Heap.bubble_up()`. Run the below code to verify the result.

<!-- grade_id: q_2_5_0, points: 0, grade: false, solution: true -->
```python
# Task 2.5 code cell

print('test pushin and popping of min or max heaps,  different item types, with the correct heap ordering')

def test_push_pop_heap():
    """create suite of heaps"""
    int_list   = [10,7,2,3,9,4,8,1,0,5,6]
    char_list  = ['p','a','h','k','q','r','l','f','o',
                    'g','s','d','m','e','v','t','b','j','w',
                                  'z','x','y','i','u','c','n']
    dupl_list  = [5.0, 2.1, 5.0, 2.2, 7.6, 4.2, 2.1, 5.0, 4.3]
    
    # create heaps ...
    int_max_heap   = Heap(hp_type=Heap.MAX_HEAP, )
    char_min_heap  = Heap(hp_type=Heap.MIN_HEAP)
    dupl_heap      = Heap() # default is: min heap
    
    for item in int_list:
        int_max_heap.push_heap(item)
    Heap.print_heap(int_max_heap, 'int max heap filled:')
    print('int max heap popped: ', end='')
    while not int_max_heap.empty():
        item = int_max_heap.pop_heap()
        print(f'{item:d}', end=' ')
    print(f'int max heap empty? {str(int_max_heap.empty()):s}')
    
    # do the same for the other heaps ...
    
    for item in char_list:
        char_min_heap.push_heap(item)
    Heap.print_heap(char_min_heap, 'char min heap filled:')
    print('char min heap popped: ', end='')
    
    ### BEGIN SOLUTION
    
    while not char_min_heap.empty():
        item = char_min_heap.pop_heap()
        print(f'{item:s}', end=' ')
    print(f'char min heap heap empty? {str(char_min_heap.empty()):s}')
    
    ### END SOLUTION
    
    
    for item in dupl_list:
        dupl_heap.push_heap(item)
    Heap.print_heap(dupl_heap, 'dupl heap filled:')
    print('dupl heap popped: ', end='')
    
    ### BEGIN SOLUTION
    
    while not dupl_heap.empty():
        item = dupl_heap.pop_heap()
        print(f'{item:.1f}', end=' ')
    print(f'dupl heap empty? {str(dupl_heap.empty()):s}')
    
    ### END SOLUTION

test_push_pop_heap()
```

Output:
```
test pushin and popping of min or max heaps,  different item types, with the correct heap ordering
int max heap filled: 10, 9, 8, 3, 7, 2, 4, 1, 0, 5, 6
int max heap popped: 10 9 8 7 6 5 4 3 2 1 0 int max heap empty? True
char min heap filled: a, b, c, f, g, d, e, k, j, q, i, h, m, l, v, t, p, o, w, z, x, y, s, u, r, n
char min heap popped: a b c d e f g h i j k l m n o p q r s t u v w x y z char min heap heap empty? True
dupl heap filled: 2.1, 2.2, 2.1, 4.3, 7.6, 5.0, 4.2, 5.0, 5.0
dupl heap popped: 2.1 2.1 2.2 4.2 4.3 5.0 5.0 5.0 7.6 dupl heap empty? True
```

<!-- grade_id: at_2_5_0, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true, assert_false

def test_pop_heap_correct():
    test_float_list = [0.01, 0.001, 0.1, 10, 100]
    test_int_list   = [10,7,2,3,9,4,8,1,0,5,6]
    test_char_list  = ['p','a','h','k','q','r','l','f','o',
                    'g','s','d','m','e','v','t','b','j','w',
                                  'z','x','y','i','u','c','n']
    test_dupl_list  = [5.0, 2.1, 5.0, 2.2, 7.6, 4.2, 2.1, 5.0, 4.3]
    
    # create heaps ...
    test_float_min_heap = Heap(hp_type=Heap.MIN_HEAP, items=test_float_list)
    test_float_max_heap = Heap(hp_type=Heap.MAX_HEAP, items=test_float_list)
    test_int_max_heap   = Heap(hp_type=Heap.MAX_HEAP, items=test_int_list)    # ...
    test_char_min_heap  = Heap(hp_type=Heap.MIN_HEAP, items=test_char_list)
    test_dupl_heap      = Heap(items=test_dupl_list) # default is: min heap
    
    # test all types ...
    for hp in (test_float_min_heap, test_float_max_heap, 
               test_int_max_heap, test_char_min_heap, test_dupl_heap):
        hp_type = hp.type
        prev = hp.pop_heap()
        while not hp.empty():
            item = hp.pop_heap()
            if hp_type == Heap.MIN_HEAP:
                assert_true( prev <= item )
            else:
                assert_true( prev >= item )
            prev = item
        
test_pop_heap_correct()

### END HIDDEN TESTS
```

#### Task 2.6
Run the below test code to verify the whole of all heap functions

<!-- grade_id: q_2_6_0, points: 0, grade: false, solution: true -->
```python
# Task 2.6 code cell

print('test all pushing and popping functions of heap')

ROOT  = 'Lars'  
NAMES = [
        'Maze', 'Leon', 'John', 'Zack', 'Chuc', 'Pete', 'Lisa', \
        'Mary', 'Rodd', 'Bert', 'Sean', 'Will', 'Dora', 'Anna' ]
ALL_NAMES = NAMES + [ROOT]

# make a heap of all these names ...

### BEGIN SOLUTION

heap = Heap(Heap.MIN_HEAP, ALL_NAMES)

### END SOLUTION

Heap.print_heap(heap, '\nheap all names :')

print('min heap popped: ', end='')
while not heap.empty():
    print(f'{heap.pop_heap():s}', end=' ')
print(f'heap empty? {str(heap.empty()):s}')

print('\ncheck, using push_heap() and pop_heap();\n')

check_heap = Heap(Heap.MIN_HEAP)

# push all names on the check heap using push_heap() ...

### BEGIN SOLUTION

for name in ALL_NAMES:
    check_heap.push_heap(name)

### END SOLUTION

Heap.print_heap(check_heap, 'check heap names :')

print('check heap popped: ', end='')
while not check_heap.empty():
    print(f'{check_heap.pop_heap():s}', end=' ')
print(f'check heap empty? {str(check_heap.empty()):s}')

```

Output:
```
test all pushing and popping functions of heap

heap all names : Anna, Bert, Dora, Mary, Chuc, Maze, John, Zack, Rodd, Leon, Sean, Will, Pete, Lisa, Lars
min heap popped: Anna Bert Chuc Dora John Lars Leon Lisa Mary Maze Pete Rodd Sean Will Zack heap empty? True

check, using push_heap() and pop_heap();

check heap names : Anna, Chuc, Bert, Mary, John, Leon, Dora, Zack, Rodd, Maze, Sean, Will, Pete, Lisa, Lars
check heap popped: Anna Bert Chuc Dora John Lars Leon Lisa Mary Maze Pete Rodd Sean Will Zack check heap empty? True
```

<!-- grade_id: at_2_6_0, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true, assert_false

def test_names_ordered_correct():
    nr   = 0 
    while not test_heap.empty():
        name = test_heap.pop_heap()
        assert_true( name == SORTED_NAMES[nr] ) 
        nr  += 1
        
ROOT  = 'Lars'  
NAMES = [
    'Maze', 'Leon', 'John', 'Zack', 'Chuc', 'Pete', 'Lisa', \
    'Mary', 'Rodd', 'Bert', 'Sean', 'Will', 'Dora', 'Anna' ]
ALL_NAMES = NAMES + [ROOT]

test_heap = Heap(Heap.MIN_HEAP, ALL_NAMES)
SORTED_NAMES = sorted(ALL_NAMES)
    
test_names_ordered_correct()

### END HIDDEN TESTS
```

## EXERCISE 3: Use `class Heap` to implement a PriorityQueue. 

A priority queue is a queue of tasks, to be handled by a server. In a priority queue, tasks enqueued and are being served according to their priority. In the below implementation, a seperate queue is maintained for each priority. Each incoming task is classified and assigned a priority `URG` (urgent), `HGH` (high), `NOR` (normal), or `LOW`, and stored in the according queue, by a Classifier. Under all circumstances, `URG`-priority taks are serviced by Servicing before `HGH`, which go before `NOR` etc. So, as an example, `NOR`-priority tasks are only serviced if no higher priority tasks are in queue anymore. 
![priority queue](./figures/prioqueue.png)

<!-- grade_id: q_3_0_0, points: 0, grade: false, solution: true -->
```python
class PriorityQueue:
    """list-based priority queue"""

    # define priority classes served in this
    # priority Q, plus labels ... (mind order)    
    URG,HGH,NOR,LOW = (0,1,2,3)
    PRIO      = (URG,HGH,NOR,LOW)
    PRIO_LBLS = ('URG','HGH','NOR','LOW')

    @staticmethod
    def get_nr_priorities_served():
        """return count of priority classes"""
        return len(PriorityQueue.PRIO)
    
    @staticmethod
    def serves(priority):
        """return T|F priority supported in this Q"""
        return priority in PriorityQueue.PRIO
    
    
    def __init__(self, nr_priorities = len(PRIO)):
        """construct a queue supporting nr_priorities queues"""
        self.__qs = [[] for p in range(nr_priorities)]
        self.__count = [0] * nr_priorities   
    

    def qsize(self, priority=None):
        """return count of queued item with prio, all if None"""
        
        ### BEGIN SOLUTION
        
        if priority is not None:
            # return a count of given prio only ...
            sm = self.__count[priority]
        else:
            # return the total
             sm = sum([self.__count[p] for p in range(len(self.PRIO))])
        
        ### END SOLUTION
        
        return sm

    
    def empty(self, priority=None):
        """return T/F is empty, for given prio, whole Q if None"""
        
        ### BEGIN SOLUTION
        
        if priority is not None:
            # is queue of given priority empty?
            is_empty = ( self.__count[priority] == 0 )
        else:
            # are all queues empty?
            is_empty = not any([self.qsize(p) for p in range(len(self.PRIO))])
        
        ### END SOLUTION
        
        return is_empty
    
    
    def enqueue(self, item, priority):
        """put item on prio queue with prio given"""
        
        ### BEGIN SOLUTION
        
        self.__qs[priority].append(item)
        self.__count[priority] += 1
        
        ### END SOLUTION

       
    def dequeue(self):
        """get the next in order task, priority from the priority queue"""
        next_task, prio = None, None
        
        ### BEGIN SOLUTION
        
        for prio in range(PriorityQueue.get_nr_priorities_served()):
            if not self.empty(priority=prio):
                next_task = self.__qs[prio].pop(-1)
                self.__count[prio] -= 1
                break  # pick the first in order to service
        
        ### END SOLUTION
        
        return next_task, prio

```

<!-- grade_id: at_3_0_0, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS 
from nose.tools import assert_true

assert_true ( 1==1 )

### END HIDDEN TESTS
```

#### Task 3.1
Finalize the code of the four-priority queue-based `class PriorityQueue` above. Then, using this class, implement the below appplication of it in a test priority queue and test the queue-ing part

<!-- grade_id: q_3_1_0, points: 0, grade: false, solution: true -->
```python
# Task 3.1 code cell

def create_queue():
    """factory function returning filled priority queue"""
    # for task indexing in TASK ...
    TASK,PRIO = (0,1)
            
    TASKS = (# TASK          PRIO
            ('Task 0', PriorityQueue.HGH),
            ('Task 1', PriorityQueue.NOR),
            ('Task 2', PriorityQueue.URG),
            ('Task 3', PriorityQueue.LOW),
            ('Task 4', PriorityQueue.URG),
            ('Task 5', PriorityQueue.LOW),
            ('Task 6', PriorityQueue.NOR),
            ('Task 7', PriorityQueue.LOW),
            ('Task 8', PriorityQueue.LOW),
            ('Task 9', PriorityQueue.NOR)
        )
    queue = PriorityQueue()
    # ... enqueue the above listed tasks ...
    for task in TASKS:
        # ... check if priority is served in this queue ...
        # ... if prio served, enque the (task, prio) ...
        # ... in the correct queue, as classified ...
        
        ### BEGIN SOLUTION
        
        if queue.serves( task[PRIO] ):
            # ... safe to store after this check ...
            queue.enqueue( task[TASK], task[PRIO] )
        
        ### END SOLUTION
        
    return queue

def printQueue(q, title=''):
    """simple console print of queue"""
    print(title)
    for prio in range(q.get_nr_priorities_served()):
        print('{:s} {:2d} tasks: '. \
              format(q.PRIO_LBLS[prio], q.qsize(prio)), end='')
        for t in range(q.qsize(prio)):
            print('T ', end='')
        print()
    print('Total tasks in queue: {:d}'. format(q.qsize()))  
        
def run_queue_ing_test():
    """demonstrate the queue-ing of class PriorityQueue"""
    
    my_queue = create_queue()
    
    # print the queue so far ..
    print('priority queue empty? {:s}'. format(str(my_queue.empty())))
    printQueue(my_queue, title='my_queue has tasks;')
    
run_queue_ing_test()
```

Output:
```
priority queue empty? False
my_queue has tasks;
URG  2 tasks: T T 
HGH  1 tasks: T 
NOR  3 tasks: T T T 
LOW  4 tasks: T T T T 
Total tasks in queue: 10
```

<!-- grade_id: at_3_1_0, points: 4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

from nose.tools import assert_false

test_pqueue = create_queue()
    
for q in range(len(PriorityQueue.PRIO)):
    # must be a task in queue q
    assert_false( test_pqueue.qsize(q) == 0 )

### END HIDDEN TESTS
```

<!-- grade_id: at_3_1_1, points: 4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

from nose.tools import assert_false

test_pqueue = create_queue()

test_TASKS = (# TASK          PRIO
        ('Task 0', PriorityQueue.HGH),
        ('Task 1', PriorityQueue.NOR),
        ('Task 2', PriorityQueue.URG),
        ('Task 3', PriorityQueue.LOW),
        ('Task 4', PriorityQueue.URG),
        ('Task 5', PriorityQueue.LOW),
        ('Task 6', PriorityQueue.NOR),
        ('Task 7', PriorityQueue.LOW),
        ('Task 8', PriorityQueue.LOW),
        ('Task 9', PriorityQueue.NOR)
    )

test_count = [0] * PriorityQueue.get_nr_priorities_served()

tsk, prio = 0, 1
for tsk in test_TASKS:
    test_count[tsk[prio]] += 1
    
assert_true( test_count[PriorityQueue.URG] == 2)
assert_true( test_count[PriorityQueue.HGH] == 1)
assert_true( test_count[PriorityQueue.NOR] == 3)
assert_true( test_count[PriorityQueue.LOW] == 4)

for q in range(PriorityQueue.get_nr_priorities_served()):
    assert_true( test_pqueue.qsize(q) == test_count[q] )

### END HIDDEN TESTS
```

<!-- grade_id: at_3_1_2, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true

test_TASKS = (# TASK          PRIO
        ('Task 0', PriorityQueue.HGH),
        ('Task 1', PriorityQueue.NOR),
        ('Task 2', PriorityQueue.URG),
        ('Task 3', PriorityQueue.LOW),
        ('Task 4', PriorityQueue.URG),
        ('Task 5', PriorityQueue.LOW),
        ('Task 6', PriorityQueue.NOR),
        ('Task 7', PriorityQueue.LOW),
        ('Task 8', PriorityQueue.LOW),
        ('Task 9', PriorityQueue.NOR)
    )
    
test_pqueue = create_queue()

assert_true( test_pqueue.qsize(None) == len(test_TASKS) )

total = 0
for q in range(PriorityQueue.get_nr_priorities_served()):
    total += test_pqueue.qsize(q)
    
assert_true( total == test_pqueue.qsize(None) )

### END HIDDEN TESTS
```

#### Task 3.2
Implement `PriorityQueue.dequeue()` in the class code above, and finalize the below test code queue o verify it.

<!-- grade_id: q_3_2_0, points: 0, grade: false, solution: true -->
```python
# Task 3.2 code cell

def run_servicing_test():
    """demonstrate the servicing of class PriorityQueue"""

    # create a filled queue ...
    my_queue = create_queue()
    
    # ... which will be served in this order ...
    urg_count = my_queue.qsize(priority=PriorityQueue.URG)
    hgh_count = my_queue.qsize(priority=PriorityQueue.HGH)
    nor_count = my_queue.qsize(priority=PriorityQueue.NOR)
    low_count = my_queue.qsize(priority=PriorityQueue.LOW)
    all_count = my_queue.qsize(priority=None)
    
    print(f'queue has {all_count} tasks in total, ' 
          f'urgent-to-low: {urg_count}, {hgh_count}, {nor_count}, {low_count} tasks in queue')
    while not my_queue.empty():
        # obtain the next task, next prio from dequeue() ...
        
        ### BEGIN SOLUTION
        
        my_next_task, my_next_prio = my_queue.dequeue()
        
        ### END SOLUTION
        
        print(f'...servicing next task: {my_next_task} with priority {PriorityQueue.PRIO_LBLS[my_next_prio]}')

run_servicing_test()
```

Output:
```
queue has 10 tasks in total, urgent-to-low: 2, 1, 3, 4 tasks in queue
...servicing next task: Task 4 with priority URG
...servicing next task: Task 2 with priority URG
...servicing next task: Task 0 with priority HGH
...servicing next task: Task 9 with priority NOR
...servicing next task: Task 6 with priority NOR
...servicing next task: Task 1 with priority NOR
...servicing next task: Task 8 with priority LOW
...servicing next task: Task 7 with priority LOW
...servicing next task: Task 5 with priority LOW
...servicing next task: Task 3 with priority LOW
```

<!-- grade_id: at_3_2_0, points: 3, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true

test_pqueue = create_queue()

test_TASKS = (# TASK          PRIO
        ('Task 0', PriorityQueue.HGH),
        ('Task 1', PriorityQueue.NOR),
        ('Task 2', PriorityQueue.URG),
        ('Task 3', PriorityQueue.LOW),
        ('Task 4', PriorityQueue.URG),
        ('Task 5', PriorityQueue.LOW),
        ('Task 6', PriorityQueue.NOR),
        ('Task 7', PriorityQueue.LOW),
        ('Task 8', PriorityQueue.LOW),
        ('Task 9', PriorityQueue.NOR)
    )
    
# ... which will be served in this order ...
urg_count = test_pqueue.qsize(priority=PriorityQueue.URG)
hgh_count = test_pqueue.qsize(priority=PriorityQueue.HGH)
nor_count = test_pqueue.qsize(priority=PriorityQueue.NOR)
low_count = test_pqueue.qsize(priority=PriorityQueue.LOW)
all_count = test_pqueue.qsize(priority=None)

for tsk in range(urg_count):
    my_next_task, my_next_prio = test_pqueue.dequeue()
    assert_true( my_next_prio == PriorityQueue.URG )
    
for tsk in range(hgh_count):
    my_next_task, my_next_prio = test_pqueue.dequeue()
    assert_true( my_next_prio == PriorityQueue.HGH )
    
for tsk in range(nor_count):
    my_next_task, my_next_prio = test_pqueue.dequeue()
    assert_true( my_next_prio == PriorityQueue.NOR )
    
for tsk in range(low_count):
    my_next_task, my_next_prio = test_pqueue.dequeue()
    assert_true( my_next_prio == PriorityQueue.LOW )
    
assert_true( test_pqueue.empty() )
    
### END HIDDEN TESTS
```

#### Task 3.3
Given the below `class ServiceTicket` of which we have implemented two **comperators**:

```python
    def __eq__(self, other):
        """return T|F self == other"""
        return self.__prio == other.__prio
```

and

```python
    def __lt__(self, other):
        """return T|F self < other"""
        return self.__prio < other.__prio
```

used in equality tests (like: a == b) and less-than tests (like: a < b). We need this for the ordering of the ServiceTickets instances. With two comperators ('=' and '<'), you can implement all other comperators, using logical operations. 

Finalize: `__le__(self, other)`, `__ge__(self, other)`, and `__gt__(self, other)` in `class ServiceTicket`, below. Then: finalize and run the test code in the Task 3.3 code cell, below

<!-- grade_id: q_3_3_0, points: 0, grade: false, solution: true -->
```python
class ServiceTicket:
    """implements a ('Task', priority) 2-tuple"""
    
    def __init__(self, task, priority):
        self.task = task
        self.prio = priority
        
    def __str__(self):
        return str((self.task, self.prio))
    
    def __eq__(self, other):
        """return T|F self == other"""
        return self.prio == other.prio
    
    def __lt__(self, other):
        """return T|F self < other"""
        return self.prio < other.prio
    
    def __le__(self, other):
        
        ### BEGIN SOLUTION
        
        return not self > other 
        
        ### END SOLUTION
        
    
    def __ge__(self, other):
        
        ### BEGIN SOLUTION
        
        return not self < other 
        
        ### END SOLUTION
        
    def __gt__(self, other):
        
        ### BEGIN SOLUTION
        
        return not ( self < other or self == other ) 
        
        ### END SOLUTION
    
```

<!-- grade_id: at_3_3_0, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true, assert_false

assert_false( ServiceTicket("task",  0) > ServiceTicket("task",  1))
assert_false( ServiceTicket("task", -1) > ServiceTicket("task",  1))
assert_false( ServiceTicket("task",  1) > ServiceTicket("task",  6))
    
### END HIDDEN TESTS
```

<!-- grade_id: at_3_3_1, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true

assert_false( ServiceTicket("task",  0) == ServiceTicket("task",  1))
assert_false( ServiceTicket("task", -1) == ServiceTicket("task",  1))
assert_false( ServiceTicket("task",  1) == ServiceTicket("task",  6))

### END HIDDEN TESTS
```

<!-- grade_id: at_3_3_2, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true

assert_false( ServiceTicket("task",  0) >= ServiceTicket("task",  1))
assert_false( ServiceTicket("task", -1) >= ServiceTicket("task",  1))
assert_false( ServiceTicket("task",  1) >= ServiceTicket("task",  6))

### END HIDDEN TESTS
```

<!-- grade_id: at_3_3_3, points: 0, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true

assert_true( ServiceTicket("task 1",  0) >= ServiceTicket("task 2",  0))
assert_true( ServiceTicket("task 1", -1) >= ServiceTicket("task 2", -2))
assert_true( ServiceTicket("task 1",  7) >= ServiceTicket("task 2",  6))

### END HIDDEN TESTS
```

<!-- grade_id: at_3_3_4, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true

assert_true( ServiceTicket("task 1", 10) > ServiceTicket("task 2",  0))
assert_true( ServiceTicket("task 1", -1) > ServiceTicket("task 2", -2))
assert_true( ServiceTicket("task 1",  7) > ServiceTicket("task 2",  6))

### END HIDDEN TESTS
```

<!-- grade_id: q_3_3_1, points: 0, grade: false, solution: true -->
```python
# Task 3.3 code cell

def demo_class_ServiceTicket():
    
    # for task indexing in TASK ...
    TASK,PRIO = (0,1)
        
    TASKS = (# TASK          PRIO
        ('Task 0', PriorityQueue.HGH),
        ('Task 1', PriorityQueue.NOR),
        ('Task 2', PriorityQueue.URG),
        ('Task 3', PriorityQueue.LOW),
        ('Task 4', PriorityQueue.URG),
        ('Task 5', PriorityQueue.LOW),
        ('Task 6', PriorityQueue.NOR),
        ('Task 7', PriorityQueue.LOW),
        ('Task 8', PriorityQueue.LOW),
        ('Task 9', PriorityQueue.NOR)
    )
    
    my_tickets = list()
    
    for tsk in TASKS:
        my_tickets.append( ServiceTicket( tsk[TASK], tsk[PRIO] ) ) 
        
    print('my ticket list;')
    for ticket in my_tickets:
        print(f'service ticket: {str(ticket):s}')
        
    # buid heaps from this list (which must make use of
    # the ServiceTicket comperator functions ...
        
    # build a min Heap from this list ...
    
    ### BEGIN SOLUTION
    
    my_min_heap = Heap(hp_type=Heap.MIN_HEAP, items=my_tickets)
    
    ### END SOLUTION
    
    # ... and a max Heap ...
    
    ### BEGIN SOLUTION
    
    my_max_heap = Heap(hp_type=Heap.MAX_HEAP, items=my_tickets)
    
    ### END SOLUTION
    
    print('\nServiceTicket min heap popped;')
    while not my_min_heap.empty():
        print(f'{str(my_min_heap.pop_heap()):s}')
        
    print('\nServiceTicket max heap popped;')
    while not my_max_heap.empty():
        print(f'{str(my_max_heap.pop_heap()):s}')
    
demo_class_ServiceTicket()
```

Output:
```
my ticket list;
service ticket: ('Task 0', 1)
service ticket: ('Task 1', 2)
service ticket: ('Task 2', 0)
service ticket: ('Task 3', 3)
service ticket: ('Task 4', 0)
service ticket: ('Task 5', 3)
service ticket: ('Task 6', 2)
service ticket: ('Task 7', 3)
service ticket: ('Task 8', 3)
service ticket: ('Task 9', 2)

ServiceTicket min heap popped;
('Task 4', 0)
('Task 2', 0)
('Task 0', 1)
('Task 1', 2)
('Task 6', 2)
('Task 9', 2)
('Task 8', 3)
('Task 3', 3)
('Task 5', 3)
('Task 7', 3)

ServiceTicket max heap popped;
('Task 3', 3)
('Task 7', 3)
('Task 8', 3)
('Task 5', 3)
('Task 1', 2)
('Task 9', 2)
('Task 6', 2)
('Task 0', 1)
('Task 4', 0)
('Task 2', 0)
```

<!-- grade_id: at_3_3_5, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true

assert_true( isinstance(my_min_heap, Heap) )
assert_true( isinstance(my_max_heap, Heap) )

### END HIDDEN TESTS
```

<!-- grade_id: at_3_3_6, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true

assert_true( my_min_heap.size() > 0 )
assert_true( my_max_heap.size() > 0 )

### END HIDDEN TESTS
```

<!-- grade_id: at_3_3_7, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true

# for task indexing in TASK ...
TASK,PRIO = (0,1)

test_TASKS = (# TASK          PRIO
    ('Task 0', PriorityQueue.HGH),
    ('Task 1', PriorityQueue.NOR),
    ('Task 2', PriorityQueue.URG),
    ('Task 3', PriorityQueue.LOW),
    ('Task 4', PriorityQueue.URG),
    ('Task 5', PriorityQueue.LOW),
    ('Task 6', PriorityQueue.NOR),
    ('Task 7', PriorityQueue.LOW),
    ('Task 8', PriorityQueue.LOW),
    ('Task 9', PriorityQueue.NOR)
)

test_tickets = list()

for tsk in test_TASKS:
    test_tickets.append( ServiceTicket( tsk[TASK], tsk[PRIO] ) ) 

# build a min Heap from this list ...
test_min_heap = Heap(hp_type=Heap.MIN_HEAP, items=test_tickets)
test_max_heap = Heap(hp_type=Heap.MAX_HEAP, items=test_tickets)

while not test_min_heap.empty():
    next_test_ticket = test_min_heap.pop_heap()
    assert_true( isinstance( next_test_ticket, ServiceTicket ))

while not test_max_heap.empty():
    next_test_ticket = test_max_heap.pop_heap()
    assert_true( isinstance( next_test_ticket, ServiceTicket ))

### END HIDDEN TESTS
```

#### Task 3.4
We are now going to create a priority queue based on a `Heap`; see the below code template for `class PriorityQueueWithHeap`. Finalize it and run the test code in the below Task 3.4 code cell

<!-- grade_id: q_3_4_0, points: 0, grade: false, solution: true -->
```python
class PriorityQueueWithHeap:
    '''PriorityQueue serving items with arbitrary priorities'''

    def __init__(self):
        self.prioq = Heap(Heap.MIN_HEAP)
    
    def size(self):
        '''return count of queued item with prio, count all if None'''
        return self.prioq.size()
    
    def empty(self):
        '''return T|F is empty'''
        return self.prioq.empty()
    
    def enqueue(self, ticket):
        '''insert typle (prio, item) in the heap'''
        
        ### BEGIN SOLUTION
        
        self.prioq.push_heap( ticket )
        
        ### END SOLUTION
    
    def dequeue(self):
        '''get next item from the heap'''
        
        ### BEGIN SOLUTION
        
        ticket = self.prioq.pop_heap()
        
        ### END SOLUTION
        
        return ticket
```

<!-- grade_id: q_3_4_1, points: 0, grade: false, solution: true -->
```python
# Task 3.4 code cell

def run_PriorityQueueWithHeap_test():
    """demonstrate the use of class PriorityQueueWithHeap"""
    
    # for task indexing in TASK ...
    TASK,PRIO = (0,1)
            
    TASKS = (# TASK          PRIO
            ('Task 0', 1),
            ('Task 1', 2),
            ('Task 2', 0),
            ('Task 3', 6),
            ('Task 4', 0),
            ('Task 5', 3),
            ('Task 6', 4),
            ('Task 7', 5),
            ('Task 8', 2),
            ('Task 9', 9)
        )
    
    my_prio_queue = PriorityQueueWithHeap()
    
    print('my priority queue with heap empty? {:s}'. format(str(my_prio_queue.empty())))
    
    
    for tsk in TASKS:
        # ... enqueue the above listed tasks ...
        
        ### BEGIN SOLUTION
        
        my_prio_queue.enqueue( ServiceTicket( tsk[TASK], tsk[PRIO] ) )
        
        ### END SOLUTION
        
        
    while not my_prio_queue.empty():
        # service the next task ...
        next_ticket = my_prio_queue.dequeue()
        print('servicing: {:s}'. format(str(next_ticket)))
    

run_PriorityQueueWithHeap_test()
```

Output:
```
my priority queue with heap empty? True
servicing: ('Task 2', 0)
servicing: ('Task 4', 0)
servicing: ('Task 0', 1)
servicing: ('Task 8', 2)
servicing: ('Task 1', 2)
servicing: ('Task 5', 3)
servicing: ('Task 6', 4)
servicing: ('Task 7', 5)
servicing: ('Task 3', 6)
servicing: ('Task 9', 9)
```

<!-- grade_id: at_3_4_0, points: 4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_true

test_tickets = (
    ('Task 2', 0),
    ('Task 4', 0),
    ('Task 0', 1),
    ('Task 8', 2),
    ('Task 1', 2),
    ('Task 5', 3),
    ('Task 6', 4),
    ('Task 7', 5),
    ('Task 3', 6),
    ('Task 9', 9)
)

# for task indexing in TASK ...
TASK,PRIO = (0,1)

tick_TASKS = (# TASK          PRIO
        ('Task 0', 1),
        ('Task 1', 2),
        ('Task 2', 0),
        ('Task 3', 6),
        ('Task 4', 0),
        ('Task 5', 3),
        ('Task 6', 4),
        ('Task 7', 5),
        ('Task 8', 2),
        ('Task 9', 9)
    )
    
test_pq = PriorityQueueWithHeap()
 
for tsk in tick_TASKS:
    # ... enqueue the above listed tasks ...
    test_pq.enqueue( ServiceTicket( tsk[TASK], tsk[PRIO] ) )
        
order = 0
while not test_pq.empty():
    # service the next task ...
    next_ticket = test_pq.dequeue()
    assert_true( next_ticket == ServiceTicket( test_tickets[order][TASK], test_tickets[order][PRIO] ) )   
    order += 1
    
### END HIDDEN TESTS
```

## Done